---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## Ex 1.1

```yaml
type: NormalExercise
key: e5774f382d
lang: sql
xp: 100
skills: 1
```

Do some data science.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

---

## Primary Keys and Foreign Keys

```yaml
type: VideoExercise
key: 33f1d40b69
xp: 50
```

`@projector_key`
65abb579759e5531013034984211ff50
