---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## Ex 1.1

```yaml
type: NormalExercise 
lang: sql
xp: 100 
skills: 1
key: e5774f382d   
```


Do some data science.


`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 33f1d40b69   
```

`@projector_key`
65abb579759e5531013034984211ff50
